# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

```python
%matplotlib inline
from matplotlib import pyplot as plt
import numpy as np
```

```python
x='12367'
X=list(x)
X
```

Output:
```
['1', '2', '3', '6', '7']
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    X=''.join(list(x))
    Y=''.join(list(y))
    if int(x)%int(y)==0 or int(y)%int(x)==0:
        for i in range(len(x)):
            for j in range(len(x)):
                if any x[i]==y[j]:
                    return True
    return False    
    raise NotImplementedError()
```

Output:
```
File "<ipython-input-16-efa7122616ce>", line 8
    if any x[i]==y[j]:
           ^
SyntaxError: invalid syntax
```

<!-- grade_id: projecteuler52a, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: projecteuler52b, points: 6, grade: true, solution: false -->
```python
assert True # leave this cell to grade the solution
```