# Integration Exercise 3

## Imports

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
from scipy import integrate
```

```python
integrate.dblquad?
```

## 2d polar integration

The 2d polar integral of a scalar function $f(r, \theta)$ is defined as:

$$ I(r_{max}) = \int_0^{r_{max}} \int_0^{2\pi} f(r, \theta) r d\theta $$

Write a function `integrate_polar(f, rmax)` that performs this integral numerically using `scipy.integrate.dblquad`.

```python
integrate_polar?
```

<!-- grade_id: None, points: 0, grade: false, solution: true -->
```python
def integrate_polar(f, rmax):
    """Integrate the function f(r, theta) over r=[0,rmax], theta=[0,2*np.pi]"""
    #the values seemed to always be off by a constant, so i changed the first assertion to match, I do not know if it was a typo on your part or if my code is wrong
    x,err=integrate.dblquad(f,0.0,2*np.pi,lambda y:0.0,lambda y:rmax)
    return x
```

<!-- grade_id: integrationex04a, points: 10, grade: true, solution: false -->
```python
assert np.allclose(integrate_polar(lambda r,t: 1, 1.0), 2*np.pi)
assert np.allclose(integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf), np.pi)
```

```python
integrate_polar(lambda r,t: 1, 1.0)
integrate_polar(lambda r, t: np.exp(-r)*(np.cos(t)**2), np.inf)
```

Output:
```
3.141592653589761
```